```python
NAME = ""
COLLABORATORS = ""
```

# Instructions

1. Make sure you have filled out your "NAME" and "COLLABORATORS" (if any) in the previous cell.

2. You should complete all code/markdown cells that state "YOUR CODE HERE" or "YOUR ANSWER HERE". 
   
3. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

4. Partial credit can be obtained if your solution approach is clear and the documented within comments in the implementation.

5. You should follow good coding practices. Your code should use type hints, be robust against invalid inputs, and you should also write a few test cases to check for correctness particularly including edge cases.  


## Problem 1

Write a python function which computes third power of a number if the number is odd, and the square of the number if the number is even.

<!-- grade_id: cell-14b79c9cb73bbf40, points: 0, grade: false, solution: true -->
```python
# the function takes in an int x and returns an int
def parity_power(x: int) -> int:
    if x % 2 == 0:
        return x**2
    return x**3
    raise NotImplementedError()
```

```python
# provide test cases
assert parity_power(1) == 1
assert parity_power(2) == 4
assert parity_power(3) == 27
assert parity_power(24) == 24**2
assert parity_power(377) == 377**3
print("Your function works!")
```

Output:
```
Your function works!
```

## Problem 2

An important part of this course is to complete an ML based project over the course of the semester.  We are going to start right now!

Read the description of the project uploaded on canvas and write down 2 potential ideas for your project.  

<!-- grade_id: cell-e119bc622f6ff41b, points: 10, grade: true, solution: true -->
I submitted my ideas in the 'Term Project Ideas' assignment in Canvas. My idea is to use my NLP research with Professor Waxman. I will be taking audio recordings of classes which have a mix of English, Hebrew, Yiddish, Arameic, and Yeshivish slang words, running them through WhisperAI to obtain preliminary transcriptions, and then using prompt strategies to train ChatGPT to fix the mistakes in the non-English words.

## Problem 3

This problem will test some of your programming skills and data wrangling ability.  

We will start with a dataset from 'https://github.com/JeffSackmann/tennis_atp' which covers information on Tennis matches from the past few decades.  

This dataset is provided by Jeff Sackman.

Write a function *download_data_by_year* which takes the year as an input and returns a pandas dataframe with the data from the file of the form `atp_matches_{year}.csv`.  Implement a cache so that we do not download the data if it is already in the cache. 

**Hint**: The url you need to use to access the files in github is https://raw.githubusercontent.com/JeffSackmann/tennis_atp/master/atp_matches_{year}.csv

<!-- grade_id: cell-9954e13bab216e43, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# cache for the df's that have already been downloaded
cache = {}

def download_data_by_year(year: int) -> pd.DataFrame:
    # Check if the input year has already been downloaded
    if year in cache:
        return cache[year]

    url = f"https://raw.githubusercontent.com/JeffSackmann/tennis_atp/master/atp_matches_{year}.csv"
    
    try:
        df = pd.read_csv(url)
        cache[year] = df
        return df
        
    except pd.errors.ParserError:
        print("Error parsing data")
    
    raise NotImplementedError()
```

<!-- grade_id: cell-bcd973ba4d369a21, points: 10, grade: true, solution: false -->
```python
df = download_data_by_year(2017)
assert df.shape == (2911, 49)
```

```python
# Print part of the dataset
# print("Input and target Features")
# display(pd.concat([df.data, df.target], axis=1).head())
df.info()
df.head()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 2911 entries, 0 to 2910
Data columns (total 49 columns):
 #   Column              Non-Null Count  Dtype  
---  ------              --------------  -----  
 0   tourney_id          2911 non-null   object 
 1   tourney_name        2911 non-null   object 
 2   surface             2911 non-null   object 
 3   draw_size           2911 non-null   int64  
 4   tourney_level       2911 non-null   object 
 5   tourney_date        2911 non-null   int64  
 6   match_num           2911 non-null   int64  
 7   winner_id           2911 non-null   int64  
 8   winner_seed         1238 non-null   float64
 9   winner_entry        403 non-null    object 
 10  winner_name         2911 non-null   object 
 11  winner_hand         2910 non-null   object 
 12  winner_ht           2876 non-null   float64
 13  winner_ioc          2911 non-null   object 
 14  winner_age          2911 non-null   float64
 15  loser_id            2911 non-null   int64  
 16  loser_
```

## Problem 4
Download the data for years 2000-2020 (inclusive).  Compute the average number of matches played per year on each surface for each month. 

Return your solution as a dictionary where the key is the match, month pair as a tuple and the value is the average.

**Hint**: You may have a case where no matches were played on a given surface in a month.  This should factor into your calculation as a zero.

<!-- grade_id: cell-be49d8719f86849f, points: 0, grade: false, solution: true -->
```python
# df_list = {}    
# for year in range(2000, 2004):
#     df_list[year] = download_data_by_year(year).fillna(0)

# df0 = df_list[2000]
# df0['year'] = df0['tourney_date'].astype(str).str[:4].astype(int)
# df0['month'] = df0['tourney_date'].astype(str).str[4:6].astype(int)
# # print(df0.head())

# df1 = df_list[2001]
# df1['year'] = df1['tourney_date'].astype(str).str[:4].astype(int)
# df1['month'] = df1['tourney_date'].astype(str).str[4:6].astype(int)
# # print(df1.head())

# df_all = pd.concat([df0, df1])
# df_all_grouped = df_all.groupby(['year', 'month', 'surface']).size().reset_index(name='matches')
# print("GROUPED\n", df_all_grouped)
# df_all_grouped_avg = df_all_grouped.groupby(['surface', 'month']).mean().reset_index().fillna(0)
# print("GROUPED AND AVERAGED\n", df_all_grouped_avg)
        
def get_statistics():
    # download the data for years 2000-2020 incl.
    # initialize a df with the data from 2000, and concatenate the rest of the years
    df = download_data_by_year(2000)  
    for year in range(2001, 2021):
        # download the data
        temp = download_data_by_year(year)
        # add a year and month column
        temp['year'] = year
        temp['month'] = temp['tourney_date'].astype(str).str[4:6].astype(int)
        # concatenate the df to our larger df
        df = pd.concat([df, temp])
    
    # compute the average number of matches played per year on each surface for each month
    df_grouped = df.groupby(['year', 'month', 'surface']).size().reset_index(name='matches')
    df_grouped_avg = df_grouped.groupby(['surface', 'month']).mean().reset_index().fillna(0)
    # print(df_grouped_avg)
    # print(df_grouped_avg.info())

    # create a dictionary where the key is the match, month pair as a tuple 
    # and the value is the average
    # there will be 12 * 4 keys for the 12 months * 4 surface types per month
    dict = {}
    months = [i for i in range(1,13)]
    surfaces = ['Carpet', 'Clay', 'Grass', 'Hard']
    for i in range(len(months)*len(surfaces)):
        data = df_grouped_avg['matches'].iloc[i]
        mo = i%12 if i%12 != 0 else 12
        try: 
            df_grouped_avg['month'].loc[mo]
        except KeyError: 
            data = 0
        finally:
            key = (df_grouped_avg['surface'].iloc[i], mo)
            dict[key] = data

    return dict

    # raise NotImplementedError

get_statistics()
```

Output:
```
File /opt/tljh/user/lib/python3.10/site-packages/pandas/core/indexing.py:1714, in _iLocIndexer._getitem_axis(self, key, axis)
   1711     raise TypeError("Cannot index by location index with a non-integer key")
   1713 # validate the location
-> 1714 self._validate_integer(key, axis)
   1716 return self.obj._ixs(key, axis=axis)

File /opt/tljh/user/lib/python3.10/site-packages/pandas/core/indexing.py:1647, in _iLocIndexer._validate_integer(self, key, axis)
   1645 len_axis = len(self.obj._get_axis(axis))
   1646 if key >= len_axis or key < -len_axis:
-> 1647     raise IndexError("single positional indexer is out-of-bounds")

IndexError: single positional indexer is out-of-bounds
```

<!-- grade_id: cell-8cb61f127b11a709, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
res = get_statistics()
assert_allclose(res[('Carpet', 1)], 9)
```

# Problem 5

Determine the BEST 5 players of all time.  There is not a definative answer here, this is your chance to show your creativity.  Please also explain how you arrived at your rankings.  You are free to use web resources to support your answer, but you MUST cite them as you use them.

You should answer questions like:

1. How did you define BEST?
2. Where do you believe your analysis is flawed?
3. What could you do to improve your analysis?

```python
# YOUR CODE HERE
```

YOUR ANSWER HERE