Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# PS4 Project

In this project, we are going to continue learning more about scikit-learn and modeling by building another complete model. We also might learn a little NLP along the way.

```python
from sklearn.datasets import fetch_20newsgroups

bunch = fetch_20newsgroups()
```

```python
print(bunch['DESCR'])

#for i in range(len(bunch.data)):
    #print(i, "example", bunch.target[i])

```

## Part a)

For each of the following words

- the
- subject
- computer
- player
- democrat

please compute the overall percentage of the documents the word appears in a document as well as for each class the percentage of documents the word appears in.  Please ignore case in your calculation.  I have provided the words in a list in the cell below.

Please provide your answer in the following format

```
ans = {
   <word>: {
       total_perc: <perc>, 
       subject_percs: {
           <subject>: <perc>
   }
}
```

If a word does not appear in the subject, do not include it in the `subject_percs`.

```python
word_list = ['the', 'subject', 'computer', 'player', 'democrat']
```

<!-- grade_id: cell-7532f8ec92d39fc6, points: 0, grade: false, solution: true -->
```python
from collections import defaultdict

def def_value():
    return 0

def compute_percentages(bunch):
    
    d = defaultdict(def_value)
    
   # YOUR CODE HERE
    
    return d
```

```python
from numpy.testing import assert_almost_equal
def assert_almost_equal(a, b):
    assert abs(a - b) < 1e-5, f"a={a:.5f}, b={b:.5f} are not euql"
```

<!-- grade_id: cell-674e347558499603, points: 10, grade: true, solution: false -->
```python
ans = compute_percentages(bunch)

```

## Part b)

Now interpret your results with a plot which shows for each word the relative percentages for each subject.  To make the plots a bit simpler, lets just consider 4 subjects:

- rec.sport.baseball
- talk.politics.misc
- comp.graphics
- talk.rec.motorcycles

these are provided as a list in the next cell.

```python
topics = ['rec.sport.baseball', 'talk.politics.misc', 'comp.graphics', 'rec.motorcycles']
```

```python
import matplotlib.pyplot as plt
import numpy as np

d = compute_percentages(bunch)

# YOUR CODE HERE 

```

## Part c)

Build a classifier to try to identifier whether the text belongs to one of the `comp` news groups.  This is a  binary classification problem.

The features will be the number of occurances of each word in the problem.  You are recommended to use Scikit-learn logistic regression for your implementation.

As hyperparameters, play with the maximum and minimum document frequency for words. (a gridsearch could be a good thing to use to find the optimal hyperparameters!)

NOTE: max_df and min_df play with what percentage of documents it will appear in  and whether it is significant or helpful, e.g. 
```
    'max_df': [100,500,1000,2000], 
    'min_df': [0,1,2,3,5,10,100]
```

Your answer should create a fitted `Estimator` called `classifier_model` which will be used in the testing,

The Scikit estimators you need to import are in the next cell

```python
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.pipeline import Pipeline
from sklearn.linear_model import LogisticRegression

import pandas as pd
```

The datasets will be prepared in the next cell.

```python
def preProcessingStep(bunch):

    #data clean up (remove extraneous data)

    #remove extraneous data that will cause overfitting
    bunch = fetch_20newsgroups(remove=('headers', 'footers', 'quotes'))

    info = {}
    data = []
    label = []

    # YOUR CODE HERE

    #create a data frame with the data and the label
    info = {'data': data, 'labels': label}
    df = pd.DataFrame(info)
    
    return df
```

## THOUGHT PROCESS: 

    Use this cell to explain the thought process behind your implementation.


```python
# implement classifier_model
from sklearn.base import BaseEstimator
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.model_selection import GridSearchCV
from sklearn.model_selection import train_test_split
 
# YOUR CODE HERE
# preprocessing, 
# pipeline with feature transformation 
# train test splits
# model fitting, 
# accuracy metrics and reports
# inference

```

<!-- grade_id: cell-14b98ec39b2ebaf6, points: 15, grade: true, solution: false -->
```python
from sklearn.metrics import f1_score, accuracy_score
from sklearn import base

assert isinstance(classifier_model, base.BaseEstimator)
# lets make sure it runs with no errors
```