```python
NAME = ""
COLLABORATORS = ""
```

# Instructions

1. Make sure you have filled out your "NAME" and "COLLABORATORS" (if any) in the previous cell.

2. You should complete all code/markdown cells that state "YOUR CODE HERE" or "YOUR ANSWER HERE". 
   
3. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

4. Partial credit can be obtained if your solution approach is clear and the documented within comments in the implementation.

5. You should follow good coding practices. Your code should use type hints, be robust against invalid inputs, and you should also write a few test cases to check for correctness particularly including edge cases.  


# Supervised Learning

In this homework, we will build a few models from scratch and then use them to explore a real world dataset.  We are hoping to get some insight into the following topics:

1. How does the training process work?
2. What is the flow for the prediction process?
3. What does a model even look like?

Most machine learning libraries or packages will make some assumptions about the input and output format of the data.  Here we will also standardize on a format which is inspired by one such library.

We will assume that our data is structured as a dataframe with a special column called "label" for the classification target.  So for a problem with two features, it might look like:

|Feature 1|Feature 2|label|
|---|---|---|
|$x_{0,0}$|$x_{0,1}$|$y_{0}$|
|$x_{1,0}$|$x_{1,1}$|$y_{1}$|
|$\vdots$|$\vdots$|$\vdots$|
|$x_{n,0}$|$x_{n,1}$|$y_{n}$|

Please make sure that your code is passing all the `assert` statements before you move onto the next part. Do not assume however, that the `assert` statements will catch all cases, you will need to do your own testing as well.

# Problem 1  - Train a decision Tree

We start with building a decision tree classifier and implement both the training and inference.  As this is our first machine learning algorithm, we will take it slow and build it in many small steps and then put it all together.

Our goal will be to build up a decision tree model using a variant of the C4.5 algorithm. This algorithm is an extension of the ID3 algorithm and solves many of the limitations of that algorithm. 

For more materials on this algorithm, please peruse

- https://en.wikipedia.org/wiki/C4.5_algorithm


The first few parts of this algorithm will be to build up a decision tree training algorithm, the second parts are to explore what to do with it.

The algorithm proceeds as follows:

1. Start with a dataset.  
2. Check base cases
3. Choose the splitting criteria which maximizes gain ratio (i.e. which feature for which you want to split)
4. Split the dataset on this critera and recurse on each split

For simplicity, we will consider two base cases:

1. If all the samples have the same label
2. If there are no features in the dataframe (this would happen as a result of recursion)

Lets start by building up some of the helper functions which we will need as part of the algorithm.

# Part a)

First we will need to define a function to compute the entropy of a numeric pandas series. we define entropy as 

$$ \sum_i - p_i \ln_2(p_i)$$

For the case of a dataset, the probably of each element occuring can be calculated from the training data quite easily. 


For the sake of checking your initial answers please use the following dataset for all the assertions.  Please do not change this dataset or you might not be able to check your assertions.

```python
import pandas as pd
trial_df = pd.DataFrame({
    'a': [1, 2, 3, 4, 2],
    'b': ['cat', 'dog', 'cat', 'dog', 'lion'],
    'label': [True, False, True, False, True]
})
trial_df
```

```python
# import numpy so we can use the almost equal method
import numpy as np
def assert_almost_equal(a, b):
    """assert almost equal with readable error message"""
    assert np.isclose(a, b), f"{a} != {b}"
```

<!-- grade_id: cell-91a03164810e3c3e, points: 0, grade: false, solution: true -->
```python
def entropy(s: pd.Series) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a78390042976f4be, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(entropy(trial_df['a']), 1.9219280948873623)
```

## Part b)
Now we consider the gain, which you can think of as the information that you gain from making a decision for a particular variable.  Intuitively, we want to choose the variable to make a split which will give us the most information once we choose.  In order to do this, we look at the difference in entropy between the full problem and the subset of the problem which would occur if we take a particular value as a given. 

Mathematically, we can represent this as

$$ G(D, ele) = H(D) - \sum P(D|ele)H(D|ele)$$
where $H$ is entropy $D$ is the dataset and $(D|ele)$ is the dataset given a chosen element.

Please define a function `gain(df, ele, label_key)` which computes the gain of column `ele` in dataframe `df` where the label column name is given as `label_key`.

<!-- grade_id: cell-331f2f6f582a3f83, points: 0, grade: false, solution: true -->
```python
def gain(df: pd.DataFrame, ele: str, label_key: str) -> float:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7a003b2e429b3c47, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(gain(trial_df, 'b', 'label'), 0.9709505944546686)
```

## Part C)

Now we can compute the gain ratio for a dataframe.  The gain ratio is defined as 

$$ GR = \frac{G}{SI}$$

Where $SI$ is the split information defined as 
$$ -\sum_i \frac{N_i}{|N|}\log_2\left(\frac{N_i}{|N|}\right)$$
where $N_i$ is the number of occurencies of a unique value for a particular element and $|N|$ is the total number of elements in that dataset.

In other words, I have a dataset like 

|Color of Shirt|
|--------|
| red|
|blue|
|red|
|blue|
|blue|

Then the split info would be computed as 

$$ SI(\text{Color of Shirt}) = -\left[\frac{N_{red}}{|N|}\log_2\left(\frac{N_{red}}{|N|}\right) + \frac{N_{blue}}{|N|}\log_2\left(\frac{N_{blue}}{|N|}\right)\right] = - \left[\frac{2}{5}\log_2\left(\frac{2}{5}\right) + \frac{3}{5}\log_2\left(\frac{3}{5}\right)\right]\approx 0.971$$

Please feel free to leverage the following resources:

- https://machinewithdata.com/2018/07/10/how-to-calculate-gain-ratio/

Please define a function `gain_ratio(df, ele, label_key)` which computes the gain ratio of column `ele` in dataframe `df` where the label column name is given as `label_key`.

<!-- grade_id: cell-6758c9a51eeb8506, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b96d042aea36da35, points: 5, grade: true, solution: false -->
```python
assert_almost_equal(gain_ratio(trial_df, 'b', 'label'), 0.6379740263133316)
```

## Part D)

So far we have only handled categorical variables, we would like to also handle numerical variables.  What we will do is create a categorial variable out of our numeric variables by create a point $p$ and then each point is either $<=p$ or $>p$.  Once we create this point, we have a categorical feature of True, False values.  The trick here is how to compute the correct $p$.  The way we will do this is as follows:

1.  What we do is for each value of the attribute, we split the dataset into all values less than or equal to that value and all values greater than that value.  
2. We then compute the gain ratio as normal.  
3. Take the $p$ which maximizes the gain ratio and use that in our fit.

In this problem, define a `max_numeric_gain_ratio` function which computes the correct $p$.  Your function should return a tuple consisting of (max_value, number_to_split)

<!-- grade_id: cell-bd7c5f57ddb07b19, points: 0, grade: false, solution: true -->
```python
def max_numeric_gain_ratio(df: pd.DataFrame, key: str, label_key: str):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-28bc9424031b48a3, points: 5, grade: true, solution: false -->
```python
res = max_numeric_gain_ratio(trial_df, 'a','label')
assert_almost_equal(res[0],  0.4459281986214854)
assert res[1] == 3, f"{res[1]} != 3"
```

## Part E)

Lets start by implementing the simple id3 algorithm which will give us a good sense of how to do the more complex algorithm.

We will make a fake dataset of 10 observations with a label of true false for how well someone sleeps.

```python
id3_data = pd.DataFrame({
    'day_type': ['long', 'long', 'short', 'medium', 'short', 'short', 'medium', 'long', 'short', 'short'],
    'weekend': [True, False, False, False, False, False, True, True, True, False],
    'good_night_before': [True, True, False, True, True, False, False, False, False, True],
    'label': [True, False, False, False, True, False, True, False, True, False]
})
id3_data
```

Next we will define an `Id3Node` class which will be used to hold the data for our algorithm.  Each of these nodes will be the node in a tree.

```python
from dataclasses import dataclass
from typing import Dict, Optional
@dataclass
class Id3Node:
    key: str
    children: Optional[Dict[str, 'Node']] = None
```

Now we can implement the `train_id3` method.  This method should take a dataframe and return an `Id3Node` which represents the trained tree.  Please note that the `gain_ratio` is a bit unconventional to use as the splitting critera, but we are using it here as it will help us generalize into the c4.5 algorithm.

*Hint*: Since we are here only implementing the ID3 algorithm, we do not need to worry about numeric columns.  Thus, we shouldn't yet need to use our `max_numeric_gain_ratio` function from before, but rather only the `gain_ratio` function.  

It is quite likely (although not strictly required) that your algorithm be recursive in nature.  Like any recursive algorithm, play close attention to the base cases defined above.  

<!-- grade_id: cell-89579c32d1cdc0d7, points: 0, grade: false, solution: true -->
```python
def train_id3(df: pd.DataFrame) -> Id3Node:
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now we can create a `predict_id3` method which takes in a single row of a `DataFrame` (which is a `Series`) and a trained `Id3Node` (the root node) to make a prediction.

The prediction should take the row and walk down the tree at each step choosing the proper child given the row information until it gets to a leaf.

<!-- grade_id: cell-3b04a7162c97cc06, points: 5, grade: true, solution: true -->
```python
def predict_id3(row: pd.Series, node: Id3Node):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now put it all together, fill in the `fit` and `predict` methods for the `Id3Model` class.

<!-- grade_id: cell-ae496f77f77055bb, points: 5, grade: true, solution: true -->
```python
class ModelNotTrainedError(ValueError):
    """model is not trained yet"""
        
class Id3Model:
    def __init__(self):
        self.tree = None
    
    def fit(self, df):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def predict(self, df):
        # YOUR CODE HERE
        raise NotImplementedError()
```

## Part E) (BONUS)

This is a bonus part to the problem.  It is not trivial to implement, but I encourage you to try!  Partial or correct solutions to this part will be worth some extra credit.

We have already implemented the ID3 algorithm, now you can use the functions before to implement the same for the numeric types.

Now we can put it all together, implement a training algorithm which produces a trained tree. To help you, we will define a class `Node` as well as a visualize function to produce an image of the node, you can choose to use your own data structure if you so wish.

```python
from dataclasses import dataclass
from typing import Dict
@dataclass
class Node:
    key: str
    numeric: bool
    pivot: float
    children: Dict[str, 'Node']
```

<!-- grade_id: cell-287faaaf5ff85def, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part f)

Now use this algorithm to fit the wine dataset below.  Give an overview of your results and an explanation of your findings.  If you did not do the bonus problem in part e), you may use the class below or the `scikit-learn` model directly.  The class is a very simple wrapper to conform to the dataframe format we have been using.


Some questions to consider:

1. Do your results make sense?

**Note**: It may take some time to train your algorithm

```python
from sklearn.tree import DecisionTreeClassifier
class DecisionTreeModel(DecisionTreeClassifier):
    
    def fit(self, df):
        if 'label' not in df:
            raise ValueError("Label is not in df")
        X = df.drop('label', axis=1)
        self.columns_ = X.columns.tolist()        
        super().fit(X.values, df['label'].values)
        return self
    
    def predict(self, df):
        X = df[self.columns_].values if isinstance(df, pd.DataFrame) else df
        return super().predict(X)
```

```python
from sklearn.datasets import load_wine
import pandas as pd
dataset = load_wine()
df = pd.DataFrame(dataset['data'], columns=dataset['feature_names'])
df['label'] = dataset['target']
df.head()
```

# Problem 2

In this problem, we will implement the perceptron algorithm as defined in the book, below is the skeleton class we will want to implement.

Your alogorithm should take as a hyperparametr `max_iters` which is the number of times it will iterate through the training set entirely.

Please include a `get_params` method which returns a `PerceptronParams` class with the current parameters of the model.  Its your choice how you store these, but this method must return this object type as it will be used to validate your results.

For consistency with the results, please initialize all weights and biases to zero.

<!-- grade_id: cell-7cd0ee9cbbaed97b, points: 5, grade: true, solution: true -->
```python
from dataclasses import dataclass
from typing import List
@dataclass
class PerceptronParams:
    weights: List[float]
    bias: float

class Perceptron:
    def __init__(self, max_iters=20):
        # YOUR CODE HERE
        raise NotImplementedError()
        
    def get_params(self) -> PerceptronParams:
        # YOUR CODE HERE
        raise NotImplementedError()
    def fit(self, df, callback=None):
        # YOUR CODE HERE
        raise NotImplementedError()
    
    def predict(self, df):
        # YOUR CODE HERE
        raise NotImplementedError()
```

## Part 2

Now lets explore linear separability with some fake datasets.

First we will import some data which is linearly separable.  This is already done in the cell below.

```python
%matplotlib inline
linear_separable_df = pd.read_csv('linearly_separable.csv')
linear_separable_df.plot.scatter(x=0, y=1, c='label', cmap='viridis')
```

Now train a perceptron and plot the best fit line on top of the scatter plot from above.

<!-- grade_id: cell-1f907e71b21bb412, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now lets examine how the line evolves as the model is trained.  Make a plot showing how the line changes per iteration.

<!-- grade_id: cell-572d2d980ae09171, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B (Bonus - extra credit)

For extra credit, implement averaging in the perceptron model.

<!-- grade_id: cell-a17fa9016ea8266e, points: 5, grade: true, solution: true -->
```python
class AveragePerceptron(Perceptron):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Now make the same plot as in the previous section.

<!-- grade_id: cell-8651ac46df11daed, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part c

Analyze the stability of the weights as a function of iteration, If you did part b, please include it in your analysis.  Some points to consider

- How stable were they over time?
- How many iterations were appropriate?

<!-- grade_id: cell-940423dbc76c5042, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 3 conceptual

For this problem, write each answer in the cell immediately following the question which will be marked.

## Part a)

The algorithm we have used to train the decision tree is a greedy algorithm.  Explain why we use a greedy algorithm and what consequences this has on the resulting model.

<!-- grade_id: cell-caf5149447628dd7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part b)

In our perceptron algorithm, *IN YOUR OWN WORDS*, why is it important to shuffle the datasets?  

<!-- grade_id: cell-b11934c7dab18f43, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part c)

Is a decision tree guaranteed to find a globally optimal solution?  If not, what are the barriers to creating an algorithm to find the globally optimal solution.

<!-- grade_id: cell-b6a56fee4b473efe, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Problem 4 - Data Problem

Problem 4 is the real world simulation problem.  Here I will simply give you a dataset and a problem, your goal is to solve the problem and list all of your assumptions as well as your results.  This is meant to simulate many of the types of problems you may see in the future on an interview.

You will be evaluated on how well you use the techniques we have learned so far in the course, you will not need to have the best model, you will be evaluated more on how you think and explain your solution.

Here we are going to use the very common dataset, california housing.


## Part a)

First of all do some exploratory data analysis and report your findings.  The following function will get the data for you, please take it from there!

```python
from sklearn.datasets import fetch_california_housing

blob = fetch_california_housing()
df = pd.DataFrame(blob['data'], columns=blob['feature_names'])
df[blob['target_names'][0]] = blob['target']
```

## Part b)

Before modeling, its often good to start with a baseline model, something simple which will give us some indication if the model we are building is actually learning something.  Lets start by building a model which predicts that the median housing price of a block is the median of the entire dataset.  Please compute the mean squared error of this model

<!-- grade_id: cell-6dbc46a56ce14db0, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part c)

In this problem we are going to use a linear regression to explore this dataset.  We will learn more about scikit-learn in the future, here you will be provided with an object to fit a dataframe to a linear model.  You are welcome to use the `scikit-learn` object directly if you so choose.

```python
from typing import List, Optional
from sklearn.linear_model import LinearRegression

class LinearModel:
    def __init__(self, fit_intercept: bool = True):
        self.fit_intercept = fit_intercept
        self.target_variable: Optional[str] = None
        self.columns: Optional[List[str]] = None
        self.model: Optional[LinearRegression] = None
        
    def _check_fit(self):
        if self.model is None:
            raise RuntimeError("model is not yet fit")
        
    def fit(self, df: pd.DataFrame, target_variable: str) -> 'LinearModel':
        """fit a dataframe and return the coefficients and intercept
        
        Parameters
        ----------
        df: pd.DataFrame
            Input dataframe for fitting, should contain target variable
        target_variable: str
            Target variable for fitting, must be in the dataframe
        """
        self.target_variable = target_variable
        X = df.drop(target_variable, axis=1)
        y = df[target_variable].values
        self.columns = X.columns
        self.model = LinearRegression(fit_intercept=self.fit_intercept)
        self.model.fit(X.values, y)
        return self
    
    @property
    def coef(self):
        self._check_fit()
        return dict(zip(self.columns, self.model.coef_))
    
    @property
    def intercept(self) -> float:
        self._check_fit()
        return self.model.intercept_

    
    def predict(self, df: pd.DataFrame):
        """
        Make predictions using the fitted model
        
        Parameters
        ----------
        df: pd.DataFrame
            Input dataframe for prediction
        """
        self._check_fit()
        X = df[self.columns].values
        return self.model.predict(X)
```

Start off by fitting the linear regression directly to the dataset and then interpret your results.  Remove the Latitude and Longitude information for now (more on this in the next part).

<!-- grade_id: cell-71f2bc4b66decbf2, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part d)

Latitude and Longitude are different sorts of features than the rest of the features in this dataset, please explain why.

**Hint**: Try making a plot

## Part e) (Bonus)

Now we are going to use the latitude and longitude to generate some new features.  Often in an ML problem, bringing more is how we can best improve our predictive accuracy.

Instead of using directly, lets create a feature which is minimum distance from a "major" city, defined as having greater than .5 million people.

I have looked up the following wikipedia:

| City | Latitude | Longitude |
| --- | ---| ---|
|Los Angeles| 34.03| -118.15|
|San Diego|32.4254| -117.0945|
|San Francisco|37.4639| -122.2459|
|San Jose |37.2010| -121.5326|
|Sacramento|38.3454| -121.2940|

You are welcome to use any distance metric you like, however, one good one would be from the geopy library. 

```python
lat_longs = [
    (34.04, -118.15),
    (32.4254, -117.0945),
    (37.4639, -122.2459),
    (37.2010, -121.5326),
    (38.3454, -121.2940)
]
```

<!-- grade_id: cell-bd9ec73649132f14, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part f)

Now please create the best linear model that you can, produce a explanation of the decision you made and why this is a "good" model.

```python

```